Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **25 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Mingyue Tang"
COLLABORATORS = "geeksforgeeks.com"
```

---

# CS110 Pre-class Work - Deterministic Quicksort

## Question 1 [time estimate: 5 minutes]

Using Figure 7.1 in Cormen et al. as a model, perform manually the partition process on the following list: A = [1,5,6,2,3,8,9,4,7]. You just need to specify the followings:
1. The array after the process is done.
2. The value of $i$ after the process is done.

<!-- grade_id: cell-7aa520f8af13679b, points: 0, grade: true, solution: true -->
1. The array : 2,1,3,4,6,5,7,8,9 ( taking the last element as the pivot)
2. The value of i is 5, which is one position before the pivot (7)

## Question 2 [time estimate: 7 minutes]

Code up a Python implementation of `partition(A, p, r)`, closely follow the pseudo-code in Cormen et al., p.172. Your function should return the index of the pivot in the array.

<!-- grade_id: cell-1ceb2600756c60ff, points: 0, grade: false, solution: true -->
```python
def partition(arr,low,high): 
    i = ( low-1 )         
    pivot = arr[high]     
  
    for j in range(low , high): 
  
        if   arr[j] <= pivot: 
          
            i = i+1 
            arr[i],arr[j] = arr[j],arr[i] 
  
    arr[i+1],arr[high] = arr[high],arr[i+1] 
    return ( i+1 ) 

arr = [1,5,6,2,3,8,9,4,7]
 
print(partition(arr,0,7))
```

Output:
```
3
```

<!-- grade_id: cell-a57b60117a7b82fb, points: 1, grade: true, solution: false -->
```python
A = [1,5,6,2,3,8,9,4,7]
assert(partition(A, 0, len(A)-1)==6)
```

Output:
```
NameError                                 Traceback (most recent call last)
<ipython-input-3-e8019dee572f> in <module>
      1 A = [1,5,6,2,3,8,9,4,7]
----> 2 assert(partition(A, 0, len(A)-1)==6)

NameError: name 'partition' is not defined
```

## Question 3 [time estimate: 3 minutes]

Code up your own Python implementation of `quicksort(A, p, r)`, using `partition(A,p,r)`.

<!-- grade_id: cell-8c39ebb8cd1aa83a, points: 0, grade: false, solution: true -->
```python
def partition(A,low,high): 
    i = ( low-1 )      
    pivot = arr[high]     
  
    for j in range(low , high): 
  
        if   arr[j] <= pivot: 
          
            i = i+1 
            arr[i],arr[j] = arr[j],arr[i] 
  
    arr[i+1],arr[high] = arr[high],arr[i+1] 
    return ( i+1 ) 
  
 

 def quickSort(A,low,high): 
    if low < high: 
  
        pi = partition(arr,low,high) 
  
        quickSort(arr, low, pi-1) 
        quickSort(arr, pi+1, high) 

```

<!-- grade_id: cell-4f822430efd456ee, points: 1, grade: true, solution: false -->
```python
A = [0]
assert(quick_sort(A, 0, 0) == [0])
A = [3,1,2]
assert(quick_sort(A, 0, 2) == [1,2,3])
```

## Question 4 [time estimate: 10 minutes] 

Explain (using experimental plots) the running time of `quick_sort` when: 
1. all elements of array A have the same value (e.g., [1,1,1])?
2. array A contains distinct elements sorted in decreasing order (e.g., [5,4,2,1])?


<!-- grade_id: cell-b58035dd5fa02329, points: 0, grade: true, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```