Before you turn this problem in, make sure everything runs as expected. First, **restart the kernel** (in the menubar, select Kernel$\rightarrow$Restart) and then **run all cells** (in the menubar, select Cell$\rightarrow$Run All).

Note that this Pre-class Work is estimated to take **32 minutes**.

Make sure you fill in any place that says `YOUR CODE HERE` or "YOUR ANSWER HERE", as well as your name and collaborators below:

```python
NAME = "Mingyue Tang"
COLLABORATORS = "leetcode.com"
```

---

# CS110 Pre-class Work - Maximum subarray problem—part II

## Question 1 [time estimate: 2 minutes]
Paste in your Python implementation of the maximum subarray from the previous class in the cell below and use that to find out the value of the maximum subarray of this array: `A = [-2, -3, 4, -1, -2, 1]`

<!-- grade_id: cell-e69c8bbcd40ca242, points: 0, grade: true, solution: true -->
```python
def max_subarray(arr,n): 

    current_maximum = arr[0]   
    maximum_so_far =arr[0]  
      
    for i in range(1,n): 
        current_maximum = max(arr[i], current_maximum + arr[i]) 
        maximum_so_far = max(maximum_so_far,current_maximum) 
          
    return maximum_so_far 
  
arr = [-3, 4, 1, 2, -1, -4, 3] 
print("Maximum sum sum is" , max_subarray(arr,len(arr)))
```

Output:
```
Maximum sum sum is 7
```

## Question 2 [time estimate: 5 minutes]
Now, your friend Joe comes and appends a single extra number at the end of the array, which becomes:  `B = [-2, -3, 4, -1, -2, 1, 8]`. Do you have to re-run the entire maximum subarray again? Explain your answer. 
The subsequent questions will help you figure out an efficient algorithmic strategy to address the last question, but make sure to write your explanation above first, before answering the remaining questions.


<!-- grade_id: cell-d4b7cd845c816263, points: 0, grade: true, solution: true -->
No it does not have to run the entire algorithm again. The first element of the new array becomes the last element in the original array, which continues the original algorithm again. Each element in B array follows this process and will return a new maximum subarray. 

## Question 3 [time estimate: 5 minutes]

**Determine if the following statement is True or False and explain your answer.**
If the maximum subarray of the array A is different than the maximum subarray of the array B (questions 1 and 2), the new maximum-subarray doesn’t need to include 8 (i.e., the newly appended element). 


<!-- grade_id: cell-9377964a8756f13b, points: 0, grade: true, solution: true -->
False. 8 might not affect the maximum of current array, but when more elements are included, 8 as an element in the array will affect the maximum in that case.

## Question 4 [time estimate: 10 minutes]
Complete the Python function `incremental_max_subarray(x, mx)` in the cell below.

This [video](https://www.youtube.com/watch?v=AAgErqQmwmA&list=PLF_a-qBXTGFektoI6JUOTRL36JlvD04BR&index=4&t=0s)  might be helpful to understand the `incremental_max_subarray` problem.

<!-- grade_id: cell-0230e459f3d701f8, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(arr,n): 

    current_maximum = arr[0]   
    maximum_so_far =arr[0]  
      
    for i in range(1,n): 
        current_maximum = max(arr[i], current_maximum + arr[i]) 
        maximum_so_far = max(maximum_so_far,current_maximum) 
          
    return maximum_so_far 
  
arr = [-3, 4, 1, 2, -1, -4, 3] 
print("Maximum sum sum is" , max_subarray(arr,len(arr)))
```

Output:
```
Maximum sum sum is 7
```

<!-- grade_id: cell-9abf19e56620ffa3, points: 1, grade: true, solution: false -->
```python
B = [-2, -3, 4, -1, -2, 1, 8]
assert(incremental_max_subarray(B, 4) == 10)
assert(incremental_max_subarray(B[:1], 0) == B[0])
```

## Question 5

### Part 1 [time estimate: 5 minutes]

Now use `incremental_max_subarray(x, mx)` iteratively, starting from the first element, to calculate the maximum subarray of a given list. Come up with a few test cases to test your code.

<!-- grade_id: cell-fd96d4ccccd99fe6, points: 0, grade: false, solution: true -->
```python
def incremental_max_subarray(arr,n): 

    current_maximum = arr[0]   
    maximum_so_far =arr[0]  
      
    for i in range(1,n): 
        current_maximum = max(arr[i], current_maximum + arr[i]) 
        maximum_so_far = max(maximum_so_far,current_maximum) 
          
    return maximum_so_far 
  
arr = [-3, 4, 1, 2, -1, -4, 3, 8, 2, -3, -1] 
print("Maximum sum sum is" , max_subarray(arr,len(arr)))
```

Output:
```
Maximum sum sum is 15
```

<!-- grade_id: cell-669ad779766aa2c3, points: 1, grade: true, solution: false -->
```python
# Please ignore this cell. This cell is for us to implement the tests 
# to see if your code works properly. 
```

### Part 2 [time estimate: 5 minutes]
Is this more efficient than the divide-and-conquer approach? Explain.

<!-- grade_id: cell-cd8f0b130a7136db, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE